// Numbas version: finer_feedback_settings
{"name": "23-4/7  Winding Temperature", "extensions": ["quantities"], "custom_part_types": [{"source": {"pk": 264, "author": {"name": "William Haynes", "pk": 2530}, "edit_page": "/part_type/264/edit"}, "name": "Electrical Quantity", "short_name": "quantity-1", "description": "<p>This extension considers active , reactive and apparent power units as incompatble. &nbsp;This prevents the student from using the wrong unit.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "plain_string(siground(settings[\"correctAnswer\"],settings[\"digits\"]))", "hint": {"static": false, "value": "settings[\"hint\"]"}, "allowEmpty": {"static": true, "value": false}}, "can_be_gap": true, "can_be_step": true, "marking_script": "sig_digits:\n(settings[\"digits\"])\n\nmark:\ncorrectif(right and good_units);\napply(student_number);\napply(unit_feedback)\n\n\nstudent_unit_string:\nreplace_regex('ohms','ohm',\n replace_regex('\u00b0', ' deg',\n   replace_regex('-', ' ' ,\n     studentAnswer[len(student_number)..len(studentAnswer)])),\"i\")\n\n\nstudent_number:\nassert(studentAnswer[0]<>'.', warn('Did you mean 0' + studentAnswer + '?'));\nmatchnumber(studentAnswer,[\"plain\",\"en\"])[0]\n\nstudent_units:\ntry(qty(1, student_unit_string),msg, feedback(msg))\n\nstudent_quantity:\ntry(decimal(student_number) * student_units,msg,qty(0,''))\n\ninterpreted_answer:\nstudent_quantity\n\ncorrect_quantity:\nsettings[\"correctAnswer\"]\n\ncorrect_units:\nunits(correct_quantity)\n\nright:\nif(settings[\"exact\"], same(answers[0],answers[1]), answers[0]=answers[1])\n\nunits:\n[student_units,units(correct_quantity)]\n\ngood_units:\nif(settings[\"exact\"],\n   same(student_units,correct_units),\n   switch( //\n    student_units='Unit not recognized', \n      false,\n    answer_is_P,\n      same(student_units,qty(1,'kW')) or same(student_units, qty(1,'W'))\n    ,\n    answer_is_Q,\n      same(student_units,qty(1,'kvar')) or same(student_units, qty(1,'var'))\n    ,\n    answer_is_S,\n      same(student_units,qty(1,'kVA')) or same(student_units, qty(1,'VA'))\n    ,\n    compatible(correct_units,student_units)\n  ))\n \n\nanswers:\n[siground(student_quantity,sig_digits),siground(correct_quantity,sig_digits)]\n\nanswer_is_P:\nsame(correct_units,qty(1,'kW')) or same(correct_units, qty(1,'W'))\n\nanswer_is_Q:\nsame(correct_units,qty(1,'kvar')) or same(correct_units, qty(1,'var'))\n\nanswer_is_S:\nsame(correct_units,qty(1,'kVA')) or same(correct_units, qty(1,'VA'))\n\nunit_feedback:\nassert(! settings[\"show_hint\"],\n    assert(good_units,\n    switch(\n       settings[\"exact\"], feedback(\"Give answer in \" + units_string(correct_units) + \".\"),\n       answer_is_P , feedback(\"Did you mean 'W' or 'kW'?\"),\n       answer_is_Q , feedback(\"Did you mean 'VAR' or 'kVAR'?\"),\n       answer_is_S , feedback(\"Did you mean 'VA' or 'kVA'?\"),\n       feedback(\"Units are incorrect.\"))))", "marking_notes": [{"name": "sig_digits", "description": "", "definition": "(settings[\"digits\"])"}, {"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "correctif(right and good_units);\napply(student_number);\napply(unit_feedback)\n"}, {"name": "student_unit_string", "description": "<p>Extracts units from submitted</p>", "definition": "replace_regex('ohms','ohm',\n replace_regex('\u00b0', ' deg',\n   replace_regex('-', ' ' ,\n     studentAnswer[len(student_number)..len(studentAnswer)])),\"i\")\n"}, {"name": "student_number", "description": "<p>string?</p>", "definition": "assert(studentAnswer[0]<>'.', warn('Did you mean 0' + studentAnswer + '?'));\nmatchnumber(studentAnswer,[\"plain\",\"en\"])[0]"}, {"name": "student_units", "description": "", "definition": "try(qty(1, student_unit_string),msg, feedback(msg))"}, {"name": "student_quantity", "description": "<p>The student's answer, interpreted as a quantity.</p>", "definition": "try(decimal(student_number) * student_units,msg,qty(0,''))"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "student_quantity"}, {"name": "correct_quantity", "description": "", "definition": "settings[\"correctAnswer\"]"}, {"name": "correct_units", "description": "", "definition": "units(correct_quantity)"}, {"name": "right", "description": "<p>if setting is checked, answers must be exactly the same, otherwise compatible units are accepted.</p>", "definition": "if(settings[\"exact\"], same(answers[0],answers[1]), answers[0]=answers[1])"}, {"name": "units", "description": "", "definition": "[student_units,units(correct_quantity)]"}, {"name": "good_units", "description": "<p>If exact option is chosen, units must exactly match answer.</p>\n<p>Otherwise, this enforces the compatable units for P, Q, and S and all other quantites. &nbsp;js_quantities considers kW, kvar, and kVA compatible.</p>", "definition": "if(settings[\"exact\"],\n   same(student_units,correct_units),\n   switch( //\n    student_units='Unit not recognized', \n      false,\n    answer_is_P,\n      same(student_units,qty(1,'kW')) or same(student_units, qty(1,'W'))\n    ,\n    answer_is_Q,\n      same(student_units,qty(1,'kvar')) or same(student_units, qty(1,'var'))\n    ,\n    answer_is_S,\n      same(student_units,qty(1,'kVA')) or same(student_units, qty(1,'VA'))\n    ,\n    compatible(correct_units,student_units)\n  ))\n "}, {"name": "answers", "description": "<p>student and correct answers in an array</p>", "definition": "[siground(student_quantity,sig_digits),siground(correct_quantity,sig_digits)]"}, {"name": "answer_is_P", "description": "", "definition": "same(correct_units,qty(1,'kW')) or same(correct_units, qty(1,'W'))"}, {"name": "answer_is_Q", "description": "", "definition": "same(correct_units,qty(1,'kvar')) or same(correct_units, qty(1,'var'))"}, {"name": "answer_is_S", "description": "", "definition": "same(correct_units,qty(1,'kVA')) or same(correct_units, qty(1,'VA'))"}, {"name": "unit_feedback", "description": "", "definition": "assert(! settings[\"show_hint\"],\n    assert(good_units,\n    switch(\n       settings[\"exact\"], feedback(\"Give answer in \" + units_string(correct_units) + \".\"),\n       answer_is_P , feedback(\"Did you mean 'W' or 'kW'?\"),\n       answer_is_Q , feedback(\"Did you mean 'VAR' or 'kVAR'?\"),\n       answer_is_S , feedback(\"Did you mean 'VA' or 'kVA'?\"),\n       feedback(\"Units are incorrect.\"))))"}], "settings": [{"name": "correctAnswer", "label": "Correct answer", "help_url": "", "hint": "The expected quantity.", "input_type": "code", "default_value": "qty(0,'')", "evaluate": true}, {"name": "digits", "label": "Significant Digits", "help_url": "", "hint": "Student answer must be correct to this number of significant digits", "input_type": "code", "default_value": "3", "evaluate": true}, {"name": "exact", "label": "Exact Units Required", "help_url": "", "hint": "If checked, answer units must exactly match expected units; if unchecked compatible units will be accepted.", "input_type": "checkbox", "default_value": false}, {"name": "show_hint", "label": "Show unit hint.", "help_url": "", "hint": "When checked, unit feedback is provided.", "input_type": "checkbox", "default_value": true}, {"name": "hint", "label": "Input Hint", "help_url": "", "hint": "If given, this hint will be shown to the students next to the answer input box.", "input_type": "string", "default_value": "", "subvars": true}], "public_availability": "restricted", "published": false, "extensions": ["quantities"]}], "resources": [["question-resources/23.4a.png", "23.4a.png"]], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "23-4/7  Winding Temperature", "tags": [], "metadata": {"description": "", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "statement": "<p>Data supplied by a motor manufacturer lists the stator resistance per phase of a certain three phase, {v}, 60-Hz, {p}, delta-connected induction motor to be {r1} &nbsp;per phase at {t1}. Stator windings are copper.</p>\n<p>A DC volt-meter-ammeter test made across two terminals of the motor, as shown in the figure, indicates {disp(I)} at {disp(TV)} &nbsp;Determine the average temperature of the windings.</p>\n<p><img src=\"resources/question-resources/23.4a.png\" width=\"366\" height=\"212\"/></p>", "advice": "<p>$R_2 = \\frac{\\var{tv}}{\\var{disp(i)}} = \\var{disp(R2)} $</p>\n<p>$R_{2\\phi} = 3/2 R_2$</p>\n<p>$T_2 = \\frac{R_{2\\phi}}{R_{1\\phi}} \\times (T_1 + k) - k = \\var{disp(T2)}$</p>", "rulesets": {}, "extensions": ["quantities"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"p": {"name": "p", "group": "Ungrouped variables", "definition": "qty(random(40..75#5),'hp')", "description": "", "templateType": "anything", "can_override": false}, "v": {"name": "v", "group": "Ungrouped variables", "definition": "qty(random(240,440,600),'V')", "description": "", "templateType": "anything", "can_override": false}, "f": {"name": "f", "group": "Ungrouped variables", "definition": "qty(random(50,60),'hz')", "description": "", "templateType": "anything", "can_override": false}, "n": {"name": "n", "group": "Ungrouped variables", "definition": "qty(120*scalar(f)/random(2,4,6),'rpm')", "description": "", "templateType": "anything", "can_override": false}, "t1": {"name": "t1", "group": "Ungrouped variables", "definition": "qty(random(20,25),'degC')", "description": "", "templateType": "anything", "can_override": false}, "tv": {"name": "tv", "group": "Ungrouped variables", "definition": "qty(random(11..11.6#0.05),'V')", "description": "<p>test voltage</p>", "templateType": "anything", "can_override": false}, "T2": {"name": "T2", "group": "Ungrouped variables", "definition": "qty(random(40..100),'degC')", "description": "", "templateType": "anything", "can_override": false}, "k": {"name": "k", "group": "Ungrouped variables", "definition": "material[1]", "description": "", "templateType": "anything", "can_override": false}, "material": {"name": "material", "group": "Ungrouped variables", "definition": "[['copper', 234.5], ['aluminum', 224.1]][random(0,1)]", "description": "", "templateType": "anything", "can_override": false}, "r2": {"name": "r2", "group": "Ungrouped variables", "definition": "(scalar(T2)+k)/(scalar(T1)+k) * R1 * 2/3", "description": "<p>(T2+k)/(T1+k) * R1</p>", "templateType": "anything", "can_override": false}, "I": {"name": "I", "group": "Ungrouped variables", "definition": "TV/R2 in \"A\"", "description": "", "templateType": "anything", "can_override": false}, "r1": {"name": "r1", "group": "Ungrouped variables", "definition": "qty(random(0.1..0.3#0.001), 'ohm')", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["p", "v", "f", "n", "t1", "tv", "T2", "k", "material", "r2", "I", "r1"], "variable_groups": [], "functions": {"disp": {"parameters": [["q", "quantity"]], "type": "anything", "language": "jme", "definition": "siground(q,4)"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": false, "customName": "", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>$T_\\text{ave}$ &nbsp;= [[0]]</p>", "gaps": [{"type": "quantity-1", "useCustomName": true, "customName": "$T_\\text{ave}$", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": false, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "settings": {"correctAnswer": "disp(T2)", "digits": "3", "exact": false, "show_hint": true, "hint": ""}}], "sortAnswers": false}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}], "resources": ["question-resources/23.4a.png"]}]}], "contributors": [{"name": "William Haynes", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/2530/"}]}